# Equações diferenciais

Uma equação diferencial, de forma bastante abstrata, é dada por
$$ \begin{cases}
    y'(t)  & = F(t, y(t)) \\
    y(t_0) & = y_0.
\end{cases} $$
Uma solução é **uma função** $y(t)$ satisfazendo tanto a _condição inicial_ $y(t_0) = y_0$
quanto a _dinâmica_ $y'(t) = F(t, y(t))$.

Por questões práticas, também é comum dar um _intervalo de solução_, ou seja,
pedir que $y$ seja definida (pelo menos) num intervalo $[a,b]$ onde $t_0 \in [a,b]$.
Aqui, vamos usar $a = t_0$,
pensando "fisicamente" que a equação diferencial modela
a evolução de alguma quantidade $y$ que conhecemos agora ($t_0$)
e que queremos saber o quanto restará num momento futuro ($t_f = b$).

## Parte 1: O método de Euler

Continuando nesta linha, é razoável discretizar o problema da seguinte forma:

1. Definimos uma sequência de pontos $a = t_0 < t_1 < t_2 < \cdots < t_n = t_f = b$ no intervalo $[a,b]$.
2. Como $f$ é derivável, temos a aproximação $y(t_{i+1}) \sim y(t_i) + (t_{i+1} - t_i)y'(t_i)$.
3. A equação diferencial permide calcular $y'(t_i) = F(t_i, y(t_i))$.
4. Isso define (recursivamente) uma sequência de aproximações $y_i$ para $y(t_i)$ pela fórmula
$$ y_{i+1} = y_i + (t_{i+1} - t_i) F(t_i, y_i).$$

A sequência mais simples de ser usada é, claro, a de $n+1$ pontos igualmente espaçados no intervalo $[a,b]$.

```python
import numpy as np
import matplotlib.pyplot as plt
```

### 1. Implemente o método de Euler

Escreva a função `euler` que divide o intervalo $[t_0, t_f]$ em `nsteps` intervalos de mesmo comprimento,
e retorna
- as `nsteps+1` extremidades `t_i`,
- as `nsteps+1` aproximações `y_i` de $y$ nestes pontos.

Dica: use `append` para construir os valores de `y_i`, mas retorne um `np.array()` para simplificar as contas.

<!-- grade_id: euler_rec, points: 0, grade: false, solution: true -->
```python
def euler(F, t0, tf, y0, nsteps=100):
    # YOUR CODE HERE
    raise NotImplementedError()
    return ts, np.array(ys)
```

Resolvendo uma EDO simples: $f' = f$.

<!-- grade_id: euler_test1, points: 2, grade: true, solution: false -->
```python
def F1(t,y):
    return y

t_100, y_100 = euler(F1, 0, 1, 1)
assert len(t_100) == 101
assert t_100[0] == 0
assert t_100[-1] == 1
```

<!-- grade_id: euler_test2, points: 2, grade: true, solution: false -->
```python
t_100, y_100 = euler(F1, 0, 1, 1)
assert len(y_100) == 101
assert y_100[0] == 1
assert( abs(y_100[-1] - 2.704813829421526) < 5e-15 )
```

O que é melhor é que o método de Euler permite resolver equações diferenciais bastante (arbitrariamente!) complicadas:

<!-- grade_id: euler_test_3, points: 2, grade: true, solution: false -->
```python
def F_dificil(t, y):
    return t*np.sin(y) - t**2*np.cos(y)

_, y_dificil = euler(F_dificil, 0, 2, 1, nsteps=100)
assert( abs(y_dificil[-1] - 3.553673725510754) < 5e-15 )
```

### 2. Qualidade da aproximação

Agora, faça um gráfico da solução numérica com mil pontos $f_{1000}$ para $f' = f$,
e da solução exata $f(t) = e^t$ no intervalo $[0,1]$.
Faça também um gráfico da diferença entre estas soluções.

<!-- grade_id: euler_graph1, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Faça o gráfico agora do erro **relativo**.

<!-- grade_id: euler_graph11, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-0e8233679e6d3a4c, points: 1, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 1
```

<!-- grade_id: cell-fb4c204f97462e9a, points: 1, grade: true, solution: false -->
```python
xdata = ax.lines[0].get_xdata()
assert min(xdata) == 0.0
assert max(xdata) == 1.0

ydata = ax.lines[0].get_ydata()
assert min(ydata) > -0.001
assert max(ydata) <  0.001
```

```python
ax = None
```

### 3. Medindo erros

Suponha que conhecemos a solução exata da EDO, que chamaremos aqui de $g(t)$ para evitar confundir com $y_i$.
Mesmo que em muitos casos seja bastante razoável usar o "erro final" $g(t_n) - y_n$,
onde $g$ é a solução exata e $y_n$ a aproximação,
vamos nos concentrar no "erro total".
E, para ficar relativamente simples, vamos usar a **média dos valores absolutos das diferenças** entre $g(t_j)$ e $y_j$.
Usamos a média, e não a soma, porque ao variar o número de passos esta quantidade se mantém homogênea.

<!-- grade_id: def_erro, points: 0, grade: false, solution: true -->
```python
def erro(g, ys, ts):
    """ Calcula o erro entre os valores de ys e g(ts).  ys e ts  devem ser vetores numpy, e  g uma função. """
    # YOUR CODE HERE
    raise NotImplementedError()
```

Como o erro total diminui ao aumentar o número de passos efetuados?
Faça um gráfico.

<!-- grade_id: graf_erro, points: 1, grade: true, solution: true -->
```python
%%time
ns = np.logspace(1,5,num=21, dtype=int)
# YOUR CODE HERE
raise NotImplementedError()
ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-778ea7cac84cbd0f, points: 2, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 1

ydata = ax.lines[0].get_ydata()
assert 1e-6 < min(ydata) < 1e-4
assert 1e-2 < max(ydata) < 1e-1
```

E comente o resultado.

<!-- grade_id: coment_erro, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### 4. Sistemas de Equações

Se você usar funções "vetorizadas" para `F`,
o seu método de Euler é capaz de resolver (aproximar) até _sistemas_ de equações diferenciais.

Por exemplo, o sistema
$$ \begin{align}
v_0' & =  v_1 \\
v_1' & = -v_0
\end{align} $$
que é equivalente à equação diferencial trigonométrica básica $y'' = -y$,
pode ser simulado assim:

<!-- grade_id: euler_2d, points: 2, grade: true, solution: false -->
```python
def F2(t, v):
    return np.array([v[1],-v[0]])

ans = np.array([-1.0505597866729235, -0.0010851561020816455])
ts, vs = euler(F2, 0, np.pi, np.array([1.0,0.0]), nsteps=100)

assert( all(abs(vs[-1]/ans - 1) < 1e-15) )
```

Dependendo de como você tenha implementado `euler()`, os valores de `vs` serão todos iguais,
pois você terá construído uma lista com `nsteps+1` vezes o mesmo `array`.
Para corrigir isso, modifique sua função `euler` para criar uma cópia de cada valor de $f_i$ antes de incluir na lista.

<!-- grade_id: new_euler, points: 0, grade: false, solution: true -->
```python
def euler(F, t0, tf, v0, nsteps=100):
    # YOUR CODE HERE
    raise NotImplementedError()
    return ts, np.array(vs)
```

Para usar as coordenadas da solução separadamente, ou seja, $f_1$ e $f_2$,
temos que transpor `v`.
(Note que os métodos constróem ao longo do tempo, e não por coordenadas)

```python
ts, vs = euler(F2, 0, np.pi, np.array([1.0,0.0]), nsteps=100)
vs = np.transpose(vs)
vs[:,:10]
```

### 5. Mais gráficos

Observe que a sua solução numérica realmente está próxima da solução analítica,
fazendo um gráfico para a função e sua derivada,
tanto numéricas como analíticas.

<!-- grade_id: cos_sin_graf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-5d346a2441fa9837, points: 1, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 4
assert len(ax.legend().texts) == 4
```

<!-- grade_id: cell-faf8107be6af3b1e, points: 1, grade: true, solution: false -->
```python
xdata = ax.lines[0].get_xdata()
assert min(xdata) == 0.0
assert max(xdata) == np.pi

ydata = [l.get_ydata() for l in ax.lines]
assert -1.1 < np.min(ydata) < -1.0
assert np.max(ydata) == 1.0
```

```python
ax = None
```

Agora, aumente o intervalo de integração para $6\pi$, e use 600 passos.

<!-- grade_id: cos_sin_graf2, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

O que aconteceu?

<!-- grade_id: cell-91158401ba7356dc, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### 6. Erros multidimensionais

Voltemos ao intervalo $[0, \pi]$.
Como se comporta o erro no caso do sistema acima, conforme o número de pontos aumenta?
Qual erro é mais importante: o da derivada ($f_2$) ou o da função original ($f_1$)?
Você vê uma razão para isso?

<!-- grade_id: erros_sin_cos, points: 0, grade: false, solution: true -->
```python
%%time
ns = np.logspace(1,5,num=21, dtype=int)
# YOUR CODE HERE
raise NotImplementedError()
ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-49e83e9d5f3391f3, points: 1, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2
```

<!-- grade_id: cell-f9dae32dd31988a8, points: 1, grade: true, solution: false -->
```python
xdata = ax.lines[0].get_xdata()
assert min(xdata) == 1e1
assert max(xdata) == 1e5

ydata = [l.get_ydata() for l in ax.lines]
assert 1e-5 < np.min(ydata) < 5e-5
assert 1e-1 < np.max(ydata) < 5e-1
```

Comente abaixo

<!-- grade_id: cell-f5fffd6a5fbd6a51, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

O erro certamente depende do intervalo "de integração" da equação diferencial.
Será que a importância relativa também depende? Faça o gráfico do erro para o intervalo $[0, \pi/2]$ e conclua.

<!-- grade_id: graf_err_sin_cos_2, points: 1, grade: true, solution: true -->
```python
%%time
ns = np.logspace(1,5,num=21, dtype=int)
# YOUR CODE HERE
raise NotImplementedError()
```

Comente

<!-- grade_id: cell-488ef29693ca30b8, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Parte 2: Métodos de maior ordem para EDOs

### Equações diferenciais e Integrais

Calcular a integral $\int_a^b h(t) \, dt$ é a mesma coisa que resolver a equação diferencial
$$ \begin{cases}
  y'(t) & = h(t) \\
  f(a)  & = 0,
\end{cases} $$
pelo Teorema Fundamental do Cálculo.
Portanto, é possível implementar uma função `int_rect` usando a função `euler`.

Isso é verdade para vários métodos de integração: muitas vezes existe um método de Equações Diferenciais cujo caso particular é um método de integração quando a ED é "trivial"
(ou seja, quando a fórmula $F$ para $y'(t)$ não depende de $y(t)$).

### O método do Trapézio para EDOs: uma equação implícita

Inversamente, podemos usar métodos de integração e "generalizar" para usar em EDOs.
Por exemplo, podemos usar a aproximação do trapézio:
$$ y(h) = y(0) + \int_0^h y'(t) \, dt \sim h \cdot \frac{y'(0) + y'(h)}{2}.$$

Substituindo a EDO $f'(t) = F(t, f(t))$, isto dá uma equação **implícita** para calcular $y(h)$.

### Trapézio 1: Equação implícita

Escreva a equação implícita que a aproximação $y_1$ de $y(h)$ satisfaz, em função de $y_0$, $h$, $F$, $\ldots$.

<!-- grade_id: cell-21e5a9a9fe6bfda0, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

### Trapézio 2: resolvendo a equação implícita de 1 passo

Escreva uma função que dá o passo do método do trapézio.

Obs: O método da secante precisa de dois pontos para iniciar.
Que valor de $y$ deve estar próximo à solução?
Se você tiver que "chutar" outro valor, este provavelmente estará mais próximo ou mais distante?
Qual deles você deve usar como primeiro e qual como segundo valor na secante?

<!-- grade_id: cell-3472403326f99e4f, points: 0, grade: false, solution: true -->
```python
# Inclua aqui seu código da secante
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a2e59b3caa310f1a, points: 0, grade: false, solution: true -->
```python
def trapezoidal_step_sec(F, tcurr, ycurr, h):
    """Resolve a equação implícita do método trapezoidal pelo método da secante, partindo de (tcurr,ycurr),
    retornando uma aproximação de  y  em  tcurr+h."""
    tnext = tcurr+h
    # YOUR CODE HERE
    raise NotImplementedError()
```

### Alguns testes simples

<!-- grade_id: cell-fc0e186f42fe6f75, points: 2, grade: true, solution: false -->
```python
def G(t,y):
    return t - np.cos(y)

ynext = trapezoidal_step_sec(G, 0, 1, 0.1)
assert abs(ynext - 0.9488541465905801) < 1e-14
```

<!-- grade_id: cell-4c6e6b6cbd203ead, points: 1, grade: true, solution: false -->
```python
ynext = trapezoidal_step_sec(G, 0, 1, 0.3)
assert abs(ynext - 0.8668725901552984) < 1e-14
```

### Trapézio 3: Resolvendo a EDO

Agora, incorporamos o passo do trapézio numa função `trapezoidal_ode`, análoga à `euler`,
que resolve a EDO $y' = F(t, y(t))$ pelo método do trapézio.

<!-- grade_id: cell-53109c937e04b22b, points: 0, grade: false, solution: true -->
```python
def trapezoidal_ode(F, t0, tf, y0, nsteps=100):
    """Calcula uma solução aproximada da equação  y' = F(t,y)  pelo método do trapézio,
    dividindo o intervalo  [t0, tf]  em  nsteps  passos,
    e usando o método da secante para resolver a equação implícita."""
    # YOUR CODE HERE
    raise NotImplementedError()
    return ts, np.array(ys)
```

<!-- grade_id: cell-2636497dc2fecc38, points: 1, grade: true, solution: false -->
```python
def F1(t,y):
    return y

t_100, y_100 = trapezoidal_ode(F1, 0, 1, 1)
assert len(t_100) == 101
assert t_100[0] == 0
assert t_100[-1] == 1
```

<!-- grade_id: cell-bad899092faf0acd, points: 1, grade: true, solution: false -->
```python
t_100, y_100 = trapezoidal_ode(F1, 0, 1, 1)
assert len(y_100) == 101
assert y_100[0] == 1
assert( abs(y_100[-1] - 2.7183044812417965) < 5e-15 )
```

<!-- grade_id: cell-c13ccd1c6e326e4a, points: 1, grade: true, solution: false -->
```python
def F_dificil(t, y):
    return t*np.sin(y) - t**2*np.cos(y)

_, y_dificil = trapezoidal_ode(F_dificil, 0, 2, 1, nsteps=100)
assert( abs(y_dificil[-1] - 3.587544094834753) < 5e-15 )
```

<!-- grade_id: cell-ec92560fa7645ec4, points: 2, grade: true, solution: false -->
```python
t_100, y_100 = trapezoidal_ode(F1, 0, 1, 1)
assert np.allclose(y_100, np.exp(t_100))
assert 1e-5 < np.max(y_100 - np.exp(t_100)) < 5e-5

t_1000, y_1000 = trapezoidal_ode(F1, 0, 1, 1, nsteps=1000)
assert np.allclose(y_1000, np.exp(t_1000), rtol=1e-7)
assert 1e-7 < np.max(y_1000 - np.exp(t_1000)) < 5e-7
```

O que a caixa acima está testando?

<!-- grade_id: cell-565ac12e85799498, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### O método do ponto médio para integrais

Esse é um método de "duplo cálculo", e explícito:

- Andamos metade do tempo $h$, usando um passo "de Euler"
- Recalculamos o passo de Euler, mas usamos a derivada "no meio do caminho",
  que é uma aproximação melhor (porque simétrica).

Em fórmulas:
- Passo explícito, $m_i = y_i + F(t_i, y_i) \cdot h/2$
- 2º passo explícito: $y_{i+1} = y_i + F(t_i + h/2, m_i) \cdot h$

### Ponto médio 1:

Escreva uma função que dá o novo ponto usando a regra do ponto médio

<!-- grade_id: cell-a5d222acf3fcf39c, points: 0, grade: false, solution: true -->
```python
def midpoint_step(F, t, y, h):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-02935fb5b443e376, points: 1, grade: true, solution: false -->
```python
def G(t,y):
    return t - np.cos(y)

ynext = midpoint_step(G, 0, 1, 0.1)
assert abs(ynext - 0.9487165172210511) < 1e-14
```

<!-- grade_id: cell-4b69d6123f624e03, points: 1, grade: true, solution: false -->
```python
ynext = midpoint_step(G, 0, 1, 0.3)
assert abs(ynext - 0.8630045486909678) < 1e-14
```

### Ponto médio 2: EDO

E agora, incorpore este passo num método de resolução de EDOs

<!-- grade_id: cell-a6fdf31371abb993, points: 0, grade: false, solution: true -->
```python
def midpoint_ode(F, t0, tf, y0, nsteps=100):
    """Calcula uma solução aproximada da equação y' = F(t,y) no intervalo  [t0, tf]  usando  nsteps  passos,
    usando o passo explícito no ponto médio."""
    # YOUR CODE HERE
    raise NotImplementedError()
    return ts, np.array(ys)
```

<!-- grade_id: cell-0517a2cc6bff6ed6, points: 1, grade: true, solution: false -->
```python
def F1(t,y):
    return y

t_100, y_100 = midpoint_ode(F1, 0, 1, 1)
assert len(t_100) == 101
assert t_100[0] == 0
assert t_100[-1] == 1
```

<!-- grade_id: cell-fa1fbe523679c33f, points: 1, grade: true, solution: false -->
```python
t_100, y_100 = midpoint_ode(F1, 0, 1, 1)
assert len(y_100) == 101
assert y_100[0] == 1
assert( abs(y_100[-1] - 2.718236862559957) < 5e-15 )
```

<!-- grade_id: cell-698ec04dabd58518, points: 1, grade: true, solution: false -->
```python
def F_dificil(t, y):
    return t*np.sin(y) - t**2*np.cos(y)

_, y_dificil = midpoint_ode(F_dificil, 0, 2, 1, nsteps=100)
assert( abs(y_dificil[-1] - 3.5874427140529868) < 5e-15 )
```

<!-- grade_id: cell-aa8742504641ae45, points: 1, grade: true, solution: false -->
```python
t_100, y_100 = midpoint_ode(F1, 0, 1, 1)
assert np.allclose(y_100, np.exp(t_100), rtol=2e-5)
assert 2e-5 < np.max(abs(y_100 - np.exp(t_100))) < 10e-5

t_1000, y_1000 = midpoint_ode(F1, 0, 1, 1, nsteps=1000)
assert np.allclose(y_1000, np.exp(t_1000), rtol=2e-7)
assert 2e-7 < np.max(abs(y_1000 - np.exp(t_1000))) < 10e-7
```

### Gráficos

### Gráficos 1: Solução de EDOs

Resolva a EDO

$$y' = t - \cos(y)$$

com condição inicial $y(0) = 1$ no intervalo $[0,5]$, usando 30 pontos,
pelos 3 métodos (Euler, Trapézio, ponto médio).
Faça o gráfico das três soluções.
Não esqueça do título e da legenda!

<!-- grade_id: cell-33a3a84d777bfa52, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-ce8a6c701ef30448, points: 2, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 3
assert len(ax.legend().texts) == 3
```

```python
ax = None
```

Esta EDO não tem "solução analítica simples".
Como você faria para saber qual das curvas está "mais correta"?
Explique sua ideia abaixo, e faça contas / gráficos na caixa seguinte

<!-- grade_id: cell-5bab6216c6c3cc36, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-f3aaa5eae6c9a1cc, points: 2, grade: true, solution: true -->
```python
# Contas / gráfico
# YOUR CODE HERE
raise NotImplementedError()
```

### Gráficos 2: Velocidade de convergência

```python
methods_3 = [euler, trapezoidal_ode, midpoint_ode]
```

Faça uma função que desenha o gráfico dos erros de solução de uma EDO, dados
- a função $F$ da EDO
- o intervalo $[a,b]$
- a condição inicial $y_0$
- a solução analítica `ysol`
- o número de pontos $n$


<!-- grade_id: cell-f1f6501fa1b5f6ce, points: 0, grade: false, solution: true -->
```python
def graph_err(F, a,b, y0, ysol, n, methods=methods_3):
    # YOUR CODE HERE
    raise NotImplementedError()
    plt.legend(title="Método")
```

Uma EDO conhecida, para comparar o erro: $$y' = -y.$$

```python
def G2(t,y): return -y
def ysol2(t): return np.exp(-t)

graph_err(G2, 0,2, 1, ysol2, 30)
plt.yscale('symlog', linthresh=1e-4) # Com "symlog", não use valores absolutos
plt.grid()
ax = plt.gca()
```

<!-- grade_id: cell-bd1283a1a93cf938, points: 1, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 3
assert len(ax.legend().texts) == 3
```

<!-- grade_id: cell-833495eb66f83962, points: 1, grade: true, solution: false -->
```python
xdata = ax.lines[0].get_xdata()
assert min(xdata) == 0.0
assert max(xdata) == 2.0

ydata = [l.get_ydata() for l in ax.lines]
assert -2e-2 < np.min(ydata) < -1e-2
assert  1e-4 < np.max(ydata) <  4e-4
```

### Gráficos 3: Erro em função de $n$

Usando a função `erro` definida anteriormente, faça o gráfico do erro, para os 3 métodos,
em função de $n$, para a EDO $y' = y$, no intervalo $[0,2]$.

<!-- grade_id: cell-042badbff1e14e15, points: 0, grade: false, solution: true -->
```python
%%time
ns = np.logspace(1,5,num=21, dtype=int)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6e8eeb4188889f83, points: 2, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 3
assert len(ax.legend().texts) == 3
```

O que você observa?

<!-- grade_id: cell-2dc4aff2a203c9c5, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE